# Extra Function Exercises

## Find the largest depositor

There are two functions that you are required to implement. For the details take a look at the code and tests below.

<!-- grade_id: cell-962f306c6dc1a36c, points: 0, grade: false, solution: true -->
```python
deposits = [
    {"name": "Fred", "amount": 10},
    {"name": "Fred", "amount": 22},
    {"name": "Wilma", "amount": 18},
    {"name": "Fred", "amount": 10},
    {"name": "Shaggy", "amount": 41},
    {"name": "Fred", "amount": 8},  
]

def summary(data):
    # The data is a list of dictionaries containing deposits. This
    # function will return a dictionary where key is the name of
    # the depositor and the value is sum of all of their deposits.
    summary = {}
    # YOUR CODE HERE
    raise NotImplementedError()
    return summary

# print(summary(deposits))
# {'Fred': 50, 'Wilma': 18, 'Shaggy': 41}

def highest(data):
    highest = None
    # YOUR CODE HERE
    raise NotImplementedError()
    return highest
```

<!-- grade_id: cell-100717b5b6366d12, points: 4, grade: true, solution: false -->
```python
#
# DO NOT EDIT THIS BLOCK
#

observed = highest(summary(deposits))
expected = {'name': 'Fred', 'amount': 50}
assert observed == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"
```

## Find the smallest depositor


<!-- grade_id: cell-3b454ced2de5a22b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-31b5ae684fe71e02, points: 1, grade: true, solution: false -->
```python
#
# DO NOT EDIT THIS BLOCK
#

observed = smallest(summary(deposits))
expected = {"name": "Wilma", "amount": 18}
assert observed == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"
```

## Expanding and string manipulation

Is this exercise you must implement two functions. This time, however, one function will call the other. For details see the test below.

<!-- grade_id: cell-f12ad5e8a668bacc, points: 0, grade: false, solution: true -->
```python
# expand_range()
# YOUR CODE HERE
raise NotImplementedError()

# expand()
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7b6f2c2718258bb4, points: 5, grade: true, solution: false -->
```python
#
# DO NOT EDIT THIS BLOCK
#

# Test the expand_range function
observed = expand_range("name", "1-5")
expected = "name1,name2,name3,name4,name5"
assert observed == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"

observed = expand("name[1,2,3]")
expected = "name1,name2,name3"
assert observed == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"

observed = expand("name[1,2,4-6]")
expected = "name1,name2,name4,name5,name6"
assert observed == expected, f"\nYou returned:\n    {observed!r}\ninstead of:\n    {expected!r}"
```